# Exercise Set 2

Below are five problems, each worth 1 point. These problems are interleaved with short tutorials on Python. This assignment will be autograded to ensure a quick turnaround. After each problem, there are tests associated with your answer which will help you determine if you have solved the problem correctly. If you can run the cell after your answer and not get any errors, then you very likely have gotten the question right. If you have errors, hopefully the error will help you identify the mistake.

Note that just because you don't get errors doesn't mean that you got the question correct. I have some additional tests held back that I do not show here, though if you pass the ones shown, you will likely pass those as well.

When you are done with the assignment, you should save this notebook manually by clicking on the save button in the toolbar (the floppy disk icon). **Do not rely on autosave. Save manually!** Ensure that you have not renamed the file. **The autograder that is used to grade this notebook requires that the file be named `Exercise_II.ipynb`.** Once you save the notebook, follow the instructions in the `README.md` file to submit the assignment.

Finally, you are encouraged to add new cells as you go through the notebook and experiment. Any cell that should not be copied or deleted is marked as such. As long as you don't copy or delete the cells marked as such, then you should feel free to experiment as much as you would like with this notebook.

The goal of this set of exercises is to introduce you to some fundamental programming concepts, particularly "if" statements and "loops." The first time you see these concepts they can be overwhelming. You are encouraged to ensure that you understand each example thoroughly. If you are struggling to understand what an example does, it is often helpful to make a copy of the cell that contains the example (this can be done by going to the toolbar, selecting "Edit" then "Copy Cell" and then "Paste Cell") and modifying the example and seeing what happens when you change something. By playing around with the examples you should be able to build intuition about what is happening with the loops and if statements.

Additionally, the exercises will require that your answers are written as functions. If you need a refresher on how to write a function (i.e. code needs to be indented and the answer needs to be `return`ed), refer back to Exercise I.

The answer to the exercise should be closely related to an example in the material near the exercise. If you find yourself doing something significantly different or beyond what is presented in the examples in this exercise, you are likely on the wrong path. This exercise set is designed to introduce all of the python syntax and concepts necessary to complete each exercise, so carefully understanding the examples in the code should allow you to complete each exercise successfully.

#### Problem \#1 - 1 point

Given a list of numbers that is at least of length 2, write a function called `sum2` that returns the sum of the first 2 elements in the list. I.e.

``` python
sum2([1, 2, 3]) → 3
sum2([1, 1]) → 2
sum2([1, 1, 1, 1]) → 2
```

You may want to refer back to the section on "Lists" in Exercise 1 as you answer the above question.

The grading cell will check the function at several values. If the function does not return the expected answer, then you will see an error. When the assignment is graded, there will be additional checks, but if you do not see an error when running the grading cell, you should feel fairly confident that your solution is correct. If you get an error, then your solution is incorrect.

<!-- grade_id: problem_1, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER SHOULD GO IN THIS CELL. DO NOT COPY THIS CELL.
def sum2(num_list):
    if len(num_list) >= 2:
        return num_list[0] + num_list[1]
    else: 
        return num_list[0]
    
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
```

<!-- grade_id: problem_1_test, points: 1, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
from nose.tools import assert_equal
assert_equal(sum2([1, 2, 3]), 3)
assert_equal(sum2([1, 1]), 2)
assert_equal(sum2([1, 1, 1, 1]), 2)
```

## Relational Operators

We have seen relational operators already in the "Intro to Python.ipynb" notebook, but they are foundational for many other aspects of programming, so we will look at them again. Fundamentally, these relational operators compare things and return either `True` or `False`.

| Symbol | Task Performed |
|----|---|
| == | True, if it is equal |
| !=  | True, if not equal to |
| < | less than |
| > | greater than |
| <=  | less than or equal to |
| >=  | greater than or equal to |

```python
z = 1
```

Note that `=` and `==` are different things. `=` sets a variable equal to something. `==` tests if two things are equal, but it doesn't set anything equal.

```python
z == 1
```

Output:
```
True
```

```python
z == 0
```

Output:
```
False
```

```python
print(z)
```

Output:
```
1
```

```python
z > 1
```

Output:
```
False
```

```python
z >= 1
```

Output:
```
True
```

Boolean operators take two boolean values (either `True` or `False`) and return a boolean values for two variables a and b, the resulting output for the `and` operator is,

| a | b | a and b |
|----|---|---|
| `True` | `True` | `True` |
| `False`  | `True` | `False` |
| `True` | `False` | `False` |
| `False` | `False` | `False` |

I.e., it's only `True` if both inputs are `True`. For the `or` operator, it is `True` if at least one of the inputs is `True`:

| a | b | a or b |
|----|---|---|
| `True` | `True` | `True` |
| `False`  | `True` | `True` |
| `True` | `False` | `True` |
| `False` | `False` | `False` |

A simple but important boolean operator is `not`,

| a | not a |
|----|---|
| `True` | `False` |
| `False`  | `True` |

```python
a = (1 > 3)
b = (3 == 3)
print("a is " + str(a))
print("b is " + str(b))
print("a or b is " + str(a or b))
print("a and b is " + str(a and b))
print("not a is " + str(not a))
```

Output:
```
a is False
b is True
a or b is True
a and b is False
not a is True
```

Why did I use `str` when I wanted to print things out? Remember, we can add strings together. Let's see what happens when we don't use it.

```python
print("a is " + a)
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[10], line 1
----> 1 print("a is " + a)

TypeError: can only concatenate str (not "bool") to str
```

This is another example of a "traceback error". This one says `TypeError: can only concatenate str (not "bool") to str`. That is becuase we are trying to "add" two things together that aren't the same. `a` is a "bool" (or a True or False value, a.k.a. Boolean), and `"a is "` is a str (or string). We can confirm that with the `type` function.

```python
print(type("a is "))
print(type(a))
```

Output:
```
<class 'str'>
<class 'bool'>
```

## `if` Statements or Making Choices

Often we want to check if a condition is True and take one action if it is, and another action if the
condition is False. We can achieve this in Python with an if statement.

__TIP:__ You can use any expression that returns a boolean value (True or False) in an if statement.

`if` statements really just give you a way to let the code make a decision. The below piece of code only prints something if `x > 0` (i.e., `x` is positive). Try changing the value of `x` to see how the result changes.

```python
x = 3

if x > 0:
    print('x is positive.')
```

Output:
```
x is positive.
```

We can make a series of decisions as well. We do this be using `if` for the first decision, and if that condition doesn't hold, we move on to the next condition. The next condition is denoted with either an `elif` (short for else if) if we want to have another check, or we can use `else` which will always run at the end. Note `else` must be the last condition checked because it will always do it's action. The below piece of code first checks to see if `x` is positive. If it is it prints `x is positive`. If it is not, it then checks to see if `x` is negative. If it is not negative either, then it does the action in the `else` statement. If it has reached the `else` statement, we can logically know that `x` must be zero since it was neither positive nor negative.

```python
# Try changing x and re-running this block, and see what happens.
x = 3
if x > 0:
    print('x is positive')
elif x < 0:
    print('x is negative')
else:
    print('x is zero')
```

Output:
```
x is positive
```

```python
# If statements can rely on boolean variables
x = -1
test = (x > 0)
print(type(test))
print(test)

if test:
    print('Test was true')
```

Output:
```
<class 'bool'>
False
```

#### Problem #2- 1 point

Write a function called `sorta_sum` that given 2 numbers, a and b, return their sum. However, sums in the range 10..19 inclusive, are forbidden, so in that case just return 20.

``` python
sorta_sum(3, 4) → 7
sorta_sum(9, 4) → 20
sorta_sum(10, 9) → 20
sorta_sum(1, 9) → 20
sorta_sum(10,11) → 21
```

You will need to handle the case when the sum is in between 10..19. This is a time where being able to do things conditionally is very helpful.

The grading cell will check the function at several values. If the function does not return the expected answer, then you will see an error. When the assignment is graded, there will be additional checks, but if you do not see an error when running the grading cell, you should feel fairly confident that your solution is correct. If you get an error, then your solution is incorrect.

<!-- grade_id: problem_2, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER SHOULD GO IN THIS CELL. DO NOT COPY THIS CELL.
def sorta_sum(a, b):
    total = a + b
    if total >= 10 and total <= 19:
        return 20
    else:
        return total
    
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
```

<!-- grade_id: problem_2_test, points: 1, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
from nose.tools import assert_equal
assert_equal(sorta_sum(3, 4), 7)
assert_equal(sorta_sum(9, 4), 20)
assert_equal(sorta_sum(10, 11), 21)
assert_equal(sorta_sum(10, 9), 20)
assert_equal(sorta_sum(1, 9), 20)
assert_equal(sorta_sum(10, 9.5), 19.5)
```

## Loops or How to Repeat Yourself

The power of programming is not in doing one thing one thing, it is in doing many things many times. Often times, we want to repeat ourselves. The two most common ways of doing this are known as `for`
loops and `while` loops. For loops in Python are useful when you want to cycle over all of the items
in a collection (such as all of the elements of an array or a list, see Exercise I for a refresher on these collections), and while loops are useful when you want to
cycle for an indefinite amount of time until some condition is met.

The basic examples below will work for looping over lists, tuples, and arrays. Looping over dictionaries
is a bit different, since there is a key and a value for each item in a dictionary. Have a look at the
Python docs for more information.

```python
# A basic for loop - don't forget the white space!
wordlist = ['hi', 'hello', 'bye']
for word in wordlist:
    print(word + '!')
```

Output:
```
hi!
hello!
bye!
```

**Note on indentation**: Notice the indentation once we enter the `for` loop.  Every idented statement after the `for` loop declaration is part of the `for` loop.  This rule holds true for `while` loops, `if` statements, functions, etc. Required identation is one of the reasons Python is such a beautiful language to read.

If you do not have consistent indentation you will get an `IndentationError`.  Fortunately, most code editors will ensure your indentation is correction.

__NOTE__ In Python the default is to use four (4) spaces for each indentation. This will happen automatically in a Jupyter Notebook if you hit tab at a place where you can indent things.

When you run the below cell, you will get an `IndentationError`. Run the cell to see the error, and then try to fix it yourself.

```python
# Indentation error: Fix it!
for word in wordlist:
    new_word = word.capitalize()
    print(new_word + '!') # Bad indent
```

Output:
```
Hi!
Hello!
Bye!
```

Here, we loop over a list and we sum up the values. Remember, even though `total = total + num` seems like a weird formula if you are thinking about high school algebra, it is perfectly fine in programming. What it is saying is take whatever is currently in the variable `total`, add `num` to it, and then overwrite the variable `total` with this new result. This is a common way to accumulate things over time.

We start with `total` being `0` so that in the first iteration of the loop, the starting value is `0`.

```python
# Sum all of the values in a collection using a for loop
numlist = [1, 4, 77, 3]
total = 0

for num in numlist:
    total = total + num
    
print("Sum is", total)
```

Output:
```
Sum is 85
```

`while` loops are useful when you don't know how many steps you will need, and want to stop once a certain condition is met. They keep going until a condition is `False`.

```python
step = 0
prod = 1
while prod < 100:
    step = step + 1
    prod = prod * 2
    print(step, prod)
    
print('Reached a product of', prod, 'at step number', step)
```

Output:
```
1 2
2 4
3 8
4 16
5 32
6 64
7 128
Reached a product of 128 at step number 7
```

We will often want to loop over a list of integers. We can do this with `range`.

```python
for i in range(10):
    print(i)
```

Output:
```
0
1
2
3
4
5
6
7
8
9
```

```python
for i in range(2,10):
    print(i)
```

Output:
```
2
3
4
5
6
7
8
9
```

```python
for i in range(2,10,4):
    print(i)
```

Output:
```
2
6
```

```python
?range
```

Output:
```
[0;31mInit signature:[0m [0mrange[0m[0;34m([0m[0mself[0m[0;34m,[0m [0;34m/[0m[0;34m,[0m [0;34m*[0m[0margs[0m[0;34m,[0m [0;34m**[0m[0mkwargs[0m[0;34m)[0m[0;34m[0m[0;34m[0m[0m
[0;31mDocstring:[0m     
range(stop) -> range object
range(start, stop[, step]) -> range object

Return an object that produces a sequence of integers from start (inclusive)
to stop (exclusive) by step.  range(i, j) produces i, i+1, i+2, ..., j-1.
start defaults to 0, and stop is omitted!  range(4) produces 0, 1, 2, 3.
These are exactly the valid indices for a list of 4 elements.
When step is given, it specifies the increment (or decrement).
[0;31mType:[0m           type
[0;31mSubclasses:[0m
```

`range` does not actually give us a `list`, it gives us a `generator` which we can pull values from. However, we can't use it like a list.

```python
a = range(2,10)
print(a)
```

Output:
```
range(2, 10)
```

A list would have printed the values. To make it a list, we can use the `list` function.

```python
print(list(a))
print(list(range(2,10)))
```

Output:
```
[2, 3, 4, 5, 6, 7, 8, 9]
[2, 3, 4, 5, 6, 7, 8, 9]
```

#### Problem \#3 - 1 point

Write a function that takes an integer `n` and computes the squares of `1,...,n`. I.e., if `n=3`, the output should be the list `[1, 4, 9]`. If `n=5`, the output should be `[1, 4, 9, 16, 25]`.

It may be helpful to start with setting a variable equal to an empty list, something like `temp = []`, and then using the `.append` method (see the "List" section of Exercise I for a refresher) to add items to the list as you loop over the integers. Also, remember that `range(1,5)` only loops from `[1,2,3,4]` (it does not include `5`). Don't forget to `return` the new list you made.

The grading cell will check the function at several values. If the function does not return the expected answer, then you will see an error. When the assignment is graded, there will be additional checks, but if you do not see an error when running the grading cell, you should feel fairly confident that your solution is correct. If you get an error, then your solution is incorrect.

<!-- grade_id: problem_3, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER SHOULD GO IN THIS CELL. DO NOT COPY THIS CELL.
def squares(n):
    """Compute squares"""
    square_list = []
    for i in range(1, n+1):
        square_list.append(i**2)
        
    return square_list
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
```

<!-- grade_id: problem_3_tests, points: 1, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
from nose.tools import assert_equal
assert_equal(squares(1), [1])
assert_equal(squares(3), [1, 4, 9])
assert_equal(squares(5), [1, 4, 9, 16, 25])
```

#### Problem \#4 - 1 point

Write a function that takes a list of numbers called `num_list`, and returns the list with only the numbers less than or equal to an optional parameter `upper_limit`. For example, if the list `[1, 7, 4, 12, 19, 5]` is the input list, and the optional parameter `upper_limit` is set to `10`, then it should return the list `[1, 7, 4, 5]`.

You will likely want to use the empty temporary list trick again for this question. However, this time you will only append to the list if the item from the `num_list` meets a certain criteria. Any time you are making a decision over things, an `if` statement is typically useful.

The grading cell will check the function at several values. If the function does not return the expected answer, then you will see an error. When the assignment is graded, there will be additional checks, but if you do not see an error when running the grading cell, you should feel fairly confident that your solution is correct. If you get an error, then your solution is incorrect.

<!-- grade_id: problem_4, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER SHOULD GO IN THIS CELL. DO NOT COPY THIS CELL.
def filter_list(num_list, upper_limit = 10):
    '''Filter a list of numbers by the upper_limit'''
    filtered_list = []
    for num in num_list:
        if num <= upper_limit:
            filtered_list.append(num)
            
    return filtered_list
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
```

<!-- grade_id: problem_4_test, points: 1, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
from nose.tools import assert_equal
assert_equal(filter_list([1, 7, 4, 12, 19, 5], upper_limit=10), [1, 7, 4, 5])
assert_equal(filter_list([1, 7, 4, 12, 19, 5]), [1, 7, 4, 5])
```

#### Problem \#5 - 1 point

Write a function called `flip_it_and_reverse_it` that takes a list called `input_list` and returns the reverse of the list. For example, the list `[1, 2, 3, 4]` would return `[4, 3, 2, 1]`. The list `["This", "is", "not", "a", "palindrome"]` would return `["palindrome", "a", "not", "is", "This"]`.

You will need to loop over the reversed list. In the "List" section of Exercise I there is a simple way to reverse a list.

The grading cell will check the function at several values. If the function does not return the expected answer, then you will see an error. When the assignment is graded, there will be additional checks, but if you do not see an error when running the grading cell, you should feel fairly confident that your solution is correct. If you get an error, then your solution is incorrect.

<!-- grade_id: problem_5, points: 0, grade: false, solution: true -->
```python
# YOUR ANSWER SHOULD GO IN THIS CELL. DO NOT COPY THIS CELL.
def flip_it_and_reverse_it(input_list):
    '''Flip the list and reverse it'''
    flipped_list = []
    for item in input_list:
        flipped_list.insert(0, item)
        
    return flipped_list
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
```

<!-- grade_id: problem_5_test, points: 1, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
from nose.tools import assert_equal
assert_equal(flip_it_and_reverse_it([1, 2, 3, 4]), [4, 3, 2, 1])
assert_equal(flip_it_and_reverse_it(["This", "is", "not", "a", "palindrome"]), ["palindrome", "a", "not", "is", "This"])
```