# Exercise Set 5

Below are five problems, each worth 1 point. These problems are interleaved with short tutorials on Python. This assignment will be autograded to ensure a quick turnaround. After several of the problems, there are tests associated with your answer which will help you determine if you have solved the problem correctly. If you can run the cell after your answer and not get any errors, then you very likely have gotten the question right. If you have errors, hopefully the error will help you identify the mistake. However, for some questions that ask for information, the grading cell will not tell you whether or not you got it correct (there is no way to do so without giving away the answer). It will tell you whether or not your formatting is correct.

Note that just because you don't get errors on the questions that do check your answer doesn't mean that you got the question correct. I have some additional tests held back that I do not show here, though if you pass the ones shown, you will likely pass those as well.

When you are done with the assignment, you should save this notebook manually by clicking on the save button in the toolbar (the floppy disk icon). **Do not rely on autosave. Save manually!** Ensure that you have not renamed the file. **The autograder that is used to grade this notebook requires that the file be named `Exercise_V.ipynb`.** Once you save the notebook, follow the instructions in the `README.md` file to submit the assignment.

Also, do not remove any datasets from the folder before downloading. You should download the folder exactly as is.

Finally, you are encouraged to add new cells as you go through the notebook and experiment. Any cell that should not be copied or deleted is marked as such. As long as you don't copy or delete the cells marked as such, then you should feel free to experiment as much as you would like with this notebook.

## Pandas DataFrames

In this exercise, we will primarily be looking at how to manipulate data within a pandas dataframe. Some of this you will have seen before in starter code, but hopefully you will develop a deeper understanding and appreciation for how to manipulate data in Python. First, we will import pandas and call it `pd` (giving us access to all of the pandas functions by writing `pd.FUNCTION`).

```python
import pandas as pd
```

I am deeply indebted to open source material available at a few places. The list (not exhaustive) of sources is the following:

  1. [Python Pandas Tutorial: A Complete Introduction for Beginners](https://www.learndatasci.com/tutorials/python-pandas-tutorial-complete-introduction-for-beginners/). A large number of the examples was copied verbatim from here.
  2. [Pandas Exercises](https://github.com/guipsamora/pandas_exercises). A GitHub repository that provides a lot of examples around using Pandas. Some of the exercises were adapted from these examples, so do not use this as a resource until after submitting this and the next exercise sets.

## Problem Dataset - Chipotle Orders

In this assignment's questions, we are going to be analyzing some data from orders at Chipotle. You can read it in with the following cell. Notice the `sep = '\t'`. This is because the data set is tab seperated, not comma separated as we are used to.

Note that the below cell drops the duplicate rows, and it creates a new column called `item_price_number`. These were transformations to the dataset that were done in Exercise III. If you would like to familiarize yourself with these steps, you are encouraged to go back to Exercise III.

```python
chipo = pd.read_csv("chipotle.tsv", sep = '\t')

chipo.drop_duplicates(inplace=True)

def price_to_float(price):
    return float(price[1:-1])

chipo['item_price_number'] = chipo['item_price'].apply(price_to_float)
```

You should spend some time familiarizing yourself with the data.

```python
chipo.head()
```

Output:
```
order_id  quantity                              item_name  \
0         1         1           Chips and Fresh Tomato Salsa   
1         1         1                                   Izze   
2         1         1                       Nantucket Nectar   
3         1         1  Chips and Tomatillo-Green Chili Salsa   
4         2         2                           Chicken Bowl   

                                  choice_description item_price  \
0                                                NaN     $2.39    
1                                       [Clementine]     $3.39    
2                                            [Apple]     $3.39    
3                                                NaN     $2.39    
4  [Tomatillo-Red Chili Salsa (Hot), [Black Beans...    $16.98    

   item_price_number  
0               2.39  
1               3.39  
2               3.39  
3               2.39  
4              16.98
```

This data set is broken up by `order_id` which corresponds to a particular order, and then within each order, there may be multiple items. The particular items are indicated by `item_name`, and any choices associated with those items are indicated by `choice_description` (e.g. the second item in `order_id=1` is an "Izze" drink with "Clementine" flavor). Additionally, there is the `item_price` for each order and the `quantity` of the item ordered. Be careful with the `item_price` as it is not the price for the item, it is the price for the item times the number of items ordered. You can see this in row 4 where `order_id=2` orders 2 chicken bowls for `$16.98` where a single chicken bowl would be half of that.

```python
chipo.isnull().sum()
```

Output:
```
order_id                 0
quantity                 0
item_name                0
choice_description    1228
item_price               0
item_price_number        0
dtype: int64
```

## Advanced Pandas Dataframes

We will be looking at movies from IMDB. This data set is on Kaggle and you can find the original [here](https://www.kaggle.com/PromptCloudHQ/imdb-data/download).

```python
movies_df = pd.read_csv("IMDB-Movie-Data.csv", index_col="Title")
```

Note that we just specified a particular column of the csv (in this case, the one called "Title") as the index.

```python
movies_df
```

Output:
```
Rank                     Genre  \
Title                                                     
Guardians of the Galaxy     1   Action,Adventure,Sci-Fi   
Prometheus                  2  Adventure,Mystery,Sci-Fi   
Split                       3           Horror,Thriller   
Sing                        4   Animation,Comedy,Family   
Suicide Squad               5  Action,Adventure,Fantasy   
...                       ...                       ...   
Secret in Their Eyes      996       Crime,Drama,Mystery   
Hostel: Part II           997                    Horror   
Step Up 2: The Streets    998       Drama,Music,Romance   
Search Party              999          Adventure,Comedy   
Nine Lives               1000     Comedy,Family,Fantasy   

                                                               Description  \
Title                                                                        
Guardians of the Galaxy  A group of intergalactic criminals are forced ...   
Prometheus             
```

What about our columns? We can see a full list with the following.

```python
movies_df.columns
```

Output:
```
Index(['Rank', 'Genre', 'Description', 'Director', 'Actors', 'Year',
       'Runtime (Minutes)', 'Rating', 'Votes', 'Revenue (Millions)',
       'Metascore'],
      dtype='object')
```

Often times it would be nice not to have spaces in the column names, so let's get rid of them.

```python
movies_df.rename(columns={
        'Runtime (Minutes)': 'Runtime', 
        'Revenue (Millions)': 'Revenue_millions'
    }, inplace=True)
```

```python
movies_df.columns
```

Output:
```
Index(['Rank', 'Genre', 'Description', 'Director', 'Actors', 'Year', 'Runtime',
       'Rating', 'Votes', 'Revenue_millions', 'Metascore'],
      dtype='object')
```

That's better, and we remembered to do `inplace=True` this time.

Sometimes we would like everything to be lower case as well, and that would be a pain to write it all out. However, we can use list comprehension to make it easy. First, let's get the lower case names.

```python
lowercase_columns = [col.lower() for col in movies_df.columns]
```

```python
lowercase_columns
```

Output:
```
['rank',
 'genre',
 'description',
 'director',
 'actors',
 'year',
 'runtime',
 'rating',
 'votes',
 'revenue_millions',
 'metascore']
```

Now, we can just set the columns of `movies_df` equal to this.

```python
movies_df.columns = lowercase_columns
```

```python
movies_df
```

Output:
```
rank                     genre  \
Title                                                     
Guardians of the Galaxy     1   Action,Adventure,Sci-Fi   
Prometheus                  2  Adventure,Mystery,Sci-Fi   
Split                       3           Horror,Thriller   
Sing                        4   Animation,Comedy,Family   
Suicide Squad               5  Action,Adventure,Fantasy   
...                       ...                       ...   
Secret in Their Eyes      996       Crime,Drama,Mystery   
Hostel: Part II           997                    Horror   
Step Up 2: The Streets    998       Drama,Music,Romance   
Search Party              999          Adventure,Comedy   
Nine Lives               1000     Comedy,Family,Fantasy   

                                                               description  \
Title                                                                        
Guardians of the Galaxy  A group of intergalactic criminals are forced ...   
Prometheus             
```

Now, they are all lowercase.

## Working with Missing Values

While we've seen how to impute values in other assignments, sometimes it's good to clean up missing values straight in the dataframe. Let's see if we have any.

```python
movies_df.isnull()
```

Output:
```
rank  genre  description  director  actors   year  \
Title                                                                         
Guardians of the Galaxy  False  False        False     False   False  False   
Prometheus               False  False        False     False   False  False   
Split                    False  False        False     False   False  False   
Sing                     False  False        False     False   False  False   
Suicide Squad            False  False        False     False   False  False   
...                        ...    ...          ...       ...     ...    ...   
Secret in Their Eyes     False  False        False     False   False  False   
Hostel: Part II          False  False        False     False   False  False   
Step Up 2: The Streets   False  False        False     False   False  False   
Search Party             False  False        False     False   False  False   
Nine Lives               False  False        False     False   False  False   
```

That's not quite the best way to look at it `.isnull()` tells us using either `True` or `False` whether any particular values is missing. We'd really like to summarize over columns. We can use `.sum()` for that. This is also an example of chaining methods together. We often end up doing this with Pandas. The first method summarizes the data in one way, and then we summarize the summary.

```python
movies_df.isnull().sum()
```

Output:
```
rank                  0
genre                 0
description           0
director              0
actors                0
year                  0
runtime               0
rating                0
votes                 0
revenue_millions    128
metascore            64
dtype: int64
```

We can find the total number of null values using a double application of the `.sum()` method, once for all of the rows in each column and then once to sum over all of the columns.

```python
movies_df.isnull().sum().sum()
```

Output:
```
192
```

### Problem #1 - 1 point

How many null values are there altogether in the Chipotle dataset? Store your answer in a variable called `null_values`. I.e, your answer should look like
``` python
null_values = x
```
where x is a numeric value. The grade will not tell you whether not you are right, but it will tell you whether your answer is in the correct format.

<!-- grade_id: problem_1, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numbers
chipotle_data = pd.read_csv('chipotle.tsv', delimiter='\t')
null_values = chipotle_data.isnull().sum().sum()
print(null_values)
assert isinstance(null_values, numbers.Number)
```

Output:
```
1246
```

<!-- grade_id: problem_1_test, points: 1, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
# This cell will not tell you whether or not your answer is right.
# If you have not assigned something to the variable null_values or you have not assigned a string,
# you will get an error, but if you just get the question wrong,
# you will not necessarily get an error.
from nose.tools import assert_equal, assert_true
import numbers
print(null_values)
assert_true(isinstance(null_values, numbers.Number))
```

Output:
```
1246
```

Now we can see that we are missing some `revenue_millions` and some `metascore` data. Let's just get rid of those rows that have missing values.

```python
movies_df.dropna()
```

Output:
```
rank                     genre  \
Title                                                      
Guardians of the Galaxy      1   Action,Adventure,Sci-Fi   
Prometheus                   2  Adventure,Mystery,Sci-Fi   
Split                        3           Horror,Thriller   
Sing                         4   Animation,Comedy,Family   
Suicide Squad                5  Action,Adventure,Fantasy   
...                        ...                       ...   
Resident Evil: Afterlife   994   Action,Adventure,Horror   
Project X                  995                    Comedy   
Hostel: Part II            997                    Horror   
Step Up 2: The Streets     998       Drama,Music,Romance   
Nine Lives                1000     Comedy,Family,Fantasy   

                                                                description  \
Title                                                                         
Guardians of the Galaxy   A group of intergalactic criminals are forced ...   
Promethe
```

We didn't do `inplace=True`, so this won't stick to our original `movies_df` dataframe, but we can tell that it got rid of 162 rows. Maybe instead of getting rid of the rows, we would have preferred to get rid of the columns that have missing data. We can do that like this.

```python
movies_df.dropna(axis=1)
```

Output:
```
rank                     genre  \
Title                                                     
Guardians of the Galaxy     1   Action,Adventure,Sci-Fi   
Prometheus                  2  Adventure,Mystery,Sci-Fi   
Split                       3           Horror,Thriller   
Sing                        4   Animation,Comedy,Family   
Suicide Squad               5  Action,Adventure,Fantasy   
...                       ...                       ...   
Secret in Their Eyes      996       Crime,Drama,Mystery   
Hostel: Part II           997                    Horror   
Step Up 2: The Streets    998       Drama,Music,Romance   
Search Party              999          Adventure,Comedy   
Nine Lives               1000     Comedy,Family,Fantasy   

                                                               description  \
Title                                                                        
Guardians of the Galaxy  A group of intergalactic criminals are forced ...   
Prometheus             
```

Remember, `axis=0` means the rows of a dataframe, and `axis=1` means the columns of a dataframe whenever we are applying things like `.dropna()`. We see that we ended up with fewer columns (we got rid of the columns `revenue_millions` and `metascore`), but have all of our rows.

Of course, we could have imputed these values as well, and we've seen how to do that with `sklearn` functions. We can do something similar in pandas, but I don't recommend it.

## Sorting the Data

Often times we want to sort the dataframe in order to answer some questions. We can do that with the `.sort_values()` method, where we tell it `by` which column.

```python
movies_df.sort_values(by = "rating")
```

Output:
```
rank                     genre  \
Title                                                     
Disaster Movie            830                    Comedy   
Don't Fuck in the Woods    43                    Horror   
Dragonball Evolution      872  Action,Adventure,Fantasy   
Tall Men                  648   Fantasy,Horror,Thriller   
Wrecker                   969    Action,Horror,Thriller   
...                       ...                       ...   
Interstellar               37    Adventure,Drama,Sci-Fi   
The Intouchables          250    Biography,Comedy,Drama   
Dangal                    118    Action,Biography,Drama   
Inception                  81   Action,Adventure,Sci-Fi   
The Dark Knight            55        Action,Crime,Drama   

                                                               description  \
Title                                                                        
Disaster Movie           Over the course of one evening, an unsuspectin...   
Don't Fuck in the Woods
```

We may want to have it sorted descending instead.

```python
movies_df.sort_values(by = "rating", ascending=False)
```

Output:
```
rank                     genre  \
Title                                                     
The Dark Knight            55        Action,Crime,Drama   
Inception                  81   Action,Adventure,Sci-Fi   
Dangal                    118    Action,Biography,Drama   
Interstellar               37    Adventure,Drama,Sci-Fi   
Kimi no na wa              97   Animation,Drama,Fantasy   
...                       ...                       ...   
Wrecker                   969    Action,Horror,Thriller   
Tall Men                  648   Fantasy,Horror,Thriller   
Dragonball Evolution      872  Action,Adventure,Fantasy   
Don't Fuck in the Woods    43                    Horror   
Disaster Movie            830                    Comedy   

                                                               description  \
Title                                                                        
The Dark Knight          When the menace known as the Joker wreaks havo...   
Inception              
```

Notice that just like before, we haven't actually changed the underlying dataframe because we didn't sort `inplace`.

```python
movies_df
```

Output:
```
rank                     genre  \
Title                                                     
Guardians of the Galaxy     1   Action,Adventure,Sci-Fi   
Prometheus                  2  Adventure,Mystery,Sci-Fi   
Split                       3           Horror,Thriller   
Sing                        4   Animation,Comedy,Family   
Suicide Squad               5  Action,Adventure,Fantasy   
...                       ...                       ...   
Secret in Their Eyes      996       Crime,Drama,Mystery   
Hostel: Part II           997                    Horror   
Step Up 2: The Streets    998       Drama,Music,Romance   
Search Party              999          Adventure,Comedy   
Nine Lives               1000     Comedy,Family,Fantasy   

                                                               description  \
Title                                                                        
Guardians of the Galaxy  A group of intergalactic criminals are forced ...   
Prometheus             
```

We can also sort by two dimensions. The following sorts first by year, and then by rating.

```python
movies_df.sort_values(by = ["year", "rating"], ascending=False)
```

Output:
```
rank                    genre  \
Title                                                                  
Dangal                                  118   Action,Biography,Drama   
Kimi no na wa                            97  Animation,Drama,Fantasy   
Koe no katachi                          862  Animation,Drama,Romance   
La La Land                                7       Comedy,Drama,Music   
Paint It Black                          479                    Drama   
...                                     ...                      ...   
Superman Returns                        925  Action,Adventure,Sci-Fi   
The Fast and the Furious: Tokyo Drift   309    Action,Crime,Thriller   
The Break-Up                            551     Comedy,Drama,Romance   
Lady in the Water                       774    Drama,Fantasy,Mystery   
Snakes on a Plane                       902   Action,Adventure,Crime   

                                                                             description  \
Title     
```

### Problem #2 - 1 point

What item was ordered the most times in a single order? Remember, the `quantity` column is the number of times an item was ordered in a single order. The grade cell will not tell you whether or not you are right. Your answer should look like
``` python
most_single_ordered = "Item Name Typed Exactly"
```
With `Item Name Typed Exactly` replaced with the exact item name in quotes (and not indented obviously). I.e., if the item ordered the most times in a single order was `Nantucket Nectar`, then your answer would look like `most_single_ordered = "Nantucket Nectar"`.

<!-- grade_id: problem_2, points: 0, grade: false, solution: true -->
```python
import pandas as pd
chipotle_data = pd.read_csv('chipotle.tsv', delimiter='\t')
item_quantity = chipotle_data.groupby('item_name')['quantity'].sum()
most_single_ordered = item_quantity.idxmax()
print(most_single_ordered)
```

Output:
```
Chicken Bowl
```

<!-- grade_id: problem_2_test, points: 1, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
# This cell will not tell you whether or not your answer is right.
# If you have not assigned something to the variable most_single_ordered or you have not assigned a string,
# you will get an error, but if you just get the question wrong,
# you will not necessarily get an error.
from nose.tools import assert_equal, assert_in
assert_equal(type(most_single_ordered), str)
print(most_single_ordered)
```

Output:
```
Chicken Bowl
```

## Summarizing the Data

There are a couple of ways we can use the DataFrame to summarize itself. In fact, we've been using these methods to build up our `summarize_dataframe` function (after this exercise, take a look at it and see if it makes more sense). The first is the `.describe()` method.

```python
movies_df.describe()
```

Output:
```
rank         year      runtime       rating         votes  \
count  1000.000000  1000.000000  1000.000000  1000.000000  1.000000e+03   
mean    500.500000  2012.783000   113.172000     6.723200  1.698083e+05   
std     288.819436     3.205962    18.810908     0.945429  1.887626e+05   
min       1.000000  2006.000000    66.000000     1.900000  6.100000e+01   
25%     250.750000  2010.000000   100.000000     6.200000  3.630900e+04   
50%     500.500000  2014.000000   111.000000     6.800000  1.107990e+05   
75%     750.250000  2016.000000   123.000000     7.400000  2.399098e+05   
max    1000.000000  2016.000000   191.000000     9.000000  1.791916e+06   

       revenue_millions   metascore  
count        872.000000  936.000000  
mean          82.956376   58.985043  
std          103.253540   17.194757  
min            0.000000   11.000000  
25%           13.270000   47.000000  
50%           47.985000   59.500000  
75%          113.715000   72.000000  
max          936.630000  100.00000
```

Note that the `.describe()` method only summarized the continous values. We can `.describe()` categorical values, but we have to call it on just the categorical values.

```python
movies_df[['genre', 'description']].describe()
```

Output:
```
genre  \
count                      1000   
unique                      207   
top     Action,Adventure,Sci-Fi   
freq                         50   

                                              description  
count                                                1000  
unique                                               1000  
top     A group of intergalactic criminals are forced ...  
freq                                                    1
```

For a categorical value, we can tell the frequency of each category with `.value_counts()`. For a really long list (like `genre` which according to the above has 207 different values) we can use `.head()` to only get the top part by chaining methods together.

```python
movies_df['genre'].value_counts().head(10)
```

Output:
```
genre
Action,Adventure,Sci-Fi       50
Drama                         48
Comedy,Drama,Romance          35
Comedy                        32
Drama,Romance                 31
Animation,Adventure,Comedy    27
Action,Adventure,Fantasy      27
Comedy,Drama                  27
Comedy,Romance                26
Crime,Drama,Thriller          24
Name: count, dtype: int64
```

For continuous variables, we can get the correlation as well. Note that we use `numeric_only=True` as part of the argument. This is necessary for operations on pandas dataframes that expect numeric values, but there is a mixture of numeric and categorical variables in the dataframe.

```python
movies_df.corr(numeric_only=True)
```

Output:
```
rank      year   runtime    rating     votes  \
rank              1.000000 -0.261605 -0.221739 -0.219555 -0.283876   
year             -0.261605  1.000000 -0.164900 -0.211219 -0.411904   
runtime          -0.221739 -0.164900  1.000000  0.392214  0.407062   
rating           -0.219555 -0.211219  0.392214  1.000000  0.511537   
votes            -0.283876 -0.411904  0.407062  0.511537  1.000000   
revenue_millions -0.271592 -0.126790  0.267953  0.217654  0.639661   
metascore        -0.191869 -0.079305  0.211978  0.631897  0.325684   

                  revenue_millions  metascore  
rank                     -0.271592  -0.191869  
year                     -0.126790  -0.079305  
runtime                   0.267953   0.211978  
rating                    0.217654   0.631897  
votes                     0.639661   0.325684  
revenue_millions          1.000000   0.142397  
metascore                 0.142397   1.000000
```

## DataFrame Slicing, Dicing, and Extracting

We regularly only want to work with a subset of the data, either a set of columns or a set of rows. Fortunately, it is easy with DataFrames to get a subset. Let's start by pulling out a single column.

```python
genre_col = movies_df['genre']
```

What did we actually get back?

```python
type(genre_col)
```

Output:
```
pandas.core.series.Series
```

```python
genre_col
```

Output:
```
Title
Guardians of the Galaxy     Action,Adventure,Sci-Fi
Prometheus                 Adventure,Mystery,Sci-Fi
Split                               Horror,Thriller
Sing                        Animation,Comedy,Family
Suicide Squad              Action,Adventure,Fantasy
                                     ...           
Secret in Their Eyes            Crime,Drama,Mystery
Hostel: Part II                              Horror
Step Up 2: The Streets          Drama,Music,Romance
Search Party                       Adventure,Comedy
Nine Lives                    Comedy,Family,Fantasy
Name: genre, Length: 1000, dtype: object
```

It's a Series! See, I told you that DataFrames were made up of Series.

If we want to get a DataFrame back, we need to give `movies_df` a _list_ of columns (it can be a list with a single item).

```python
genre_df = movies_df[['genre']]
```

```python
type(genre_df)
```

Output:
```
pandas.core.frame.DataFrame
```

```python
genre_df
```

Output:
```
genre
Title                                            
Guardians of the Galaxy   Action,Adventure,Sci-Fi
Prometheus               Adventure,Mystery,Sci-Fi
Split                             Horror,Thriller
Sing                      Animation,Comedy,Family
Suicide Squad            Action,Adventure,Fantasy
...                                           ...
Secret in Their Eyes          Crime,Drama,Mystery
Hostel: Part II                            Horror
Step Up 2: The Streets        Drama,Music,Romance
Search Party                     Adventure,Comedy
Nine Lives                  Comedy,Family,Fantasy

[1000 rows x 1 columns]
```

If we are giving a list, we can give more columns.

```python
subset_df = movies_df[['genre', 'rating']]
```

```python
subset_df
```

Output:
```
genre  rating
Title                                                    
Guardians of the Galaxy   Action,Adventure,Sci-Fi     8.1
Prometheus               Adventure,Mystery,Sci-Fi     7.0
Split                             Horror,Thriller     7.3
Sing                      Animation,Comedy,Family     7.2
Suicide Squad            Action,Adventure,Fantasy     6.2
...                                           ...     ...
Secret in Their Eyes          Crime,Drama,Mystery     6.2
Hostel: Part II                            Horror     5.5
Step Up 2: The Streets        Drama,Music,Romance     6.2
Search Party                     Adventure,Comedy     5.6
Nine Lives                  Comedy,Family,Fantasy     5.3

[1000 rows x 2 columns]
```

For rows, we have two equivalent ways to get data. Both methods **loc**ate a row and return it.
  * `.loc` - **loc**ates by name
  * `.iloc` - **loc**ates by numerical **i**ndex
  
Let's get the data for the movie Prometheus.

```python
prom = movies_df.loc["Prometheus"]
```

```python
prom
```

Output:
```
rank                                                                2
genre                                        Adventure,Mystery,Sci-Fi
description         Following clues to the origin of mankind, a te...
director                                                 Ridley Scott
actors              Noomi Rapace, Logan Marshall-Green, Michael Fa...
year                                                             2012
runtime                                                           124
rating                                                            7.0
votes                                                          485820
revenue_millions                                               126.46
metascore                                                        65.0
Name: Prometheus, dtype: object
```

Alternatively, we can give it the row number. Remember, Prometheus is the second row in the DataFrame (and we start counting from 0).

```python
prom = movies_df.iloc[1]
```

```python
prom
```

Output:
```
rank                                                                2
genre                                        Adventure,Mystery,Sci-Fi
description         Following clues to the origin of mankind, a te...
director                                                 Ridley Scott
actors              Noomi Rapace, Logan Marshall-Green, Michael Fa...
year                                                             2012
runtime                                                           124
rating                                                            7.0
votes                                                          485820
revenue_millions                                               126.46
metascore                                                        65.0
Name: Prometheus, dtype: object
```

We can also get a range of rows, just like we can with a list. Let's look at a subset of the movies.

```python
movie_subset = movies_df.loc['Prometheus':'Sing']
```

```python
movie_subset
```

Output:
```
rank                     genre  \
Title                                        
Prometheus     2  Adventure,Mystery,Sci-Fi   
Split          3           Horror,Thriller   
Sing           4   Animation,Comedy,Family   

                                                  description  \
Title                                                           
Prometheus  Following clues to the origin of mankind, a te...   
Split       Three girls are kidnapped by a man with a diag...   
Sing        In a city of humanoid animals, a hustling thea...   

                        director  \
Title                              
Prometheus          Ridley Scott   
Split         M. Night Shyamalan   
Sing        Christophe Lourdelet   

                                                       actors  year  runtime  \
Title                                                                          
Prometheus  Noomi Rapace, Logan Marshall-Green, Michael Fa...  2012      124   
Split       James McAvoy, Anya Tay
```

```python
movie_subset = movies_df.iloc[1:4]
```

```python
movie_subset
```

Output:
```
rank                     genre  \
Title                                        
Prometheus     2  Adventure,Mystery,Sci-Fi   
Split          3           Horror,Thriller   
Sing           4   Animation,Comedy,Family   

                                                  description  \
Title                                                           
Prometheus  Following clues to the origin of mankind, a te...   
Split       Three girls are kidnapped by a man with a diag...   
Sing        In a city of humanoid animals, a hustling thea...   

                        director  \
Title                              
Prometheus          Ridley Scott   
Split         M. Night Shyamalan   
Sing        Christophe Lourdelet   

                                                       actors  year  runtime  \
Title                                                                          
Prometheus  Noomi Rapace, Logan Marshall-Green, Michael Fa...  2012      124   
Split       James McAvoy, Anya Tay
```

We can also use `.loc` and `.iloc` to access specific columns as well. The way we do this is to use the "second" slot in either `.loc` or `.iloc`.

```python
column_subset = movies_df.loc[:,"genre"]
```

```python
column_subset
```

Output:
```
Title
Guardians of the Galaxy     Action,Adventure,Sci-Fi
Prometheus                 Adventure,Mystery,Sci-Fi
Split                               Horror,Thriller
Sing                        Animation,Comedy,Family
Suicide Squad              Action,Adventure,Fantasy
                                     ...           
Secret in Their Eyes            Crime,Drama,Mystery
Hostel: Part II                              Horror
Step Up 2: The Streets          Drama,Music,Romance
Search Party                       Adventure,Comedy
Nine Lives                    Comedy,Family,Fantasy
Name: genre, Length: 1000, dtype: object
```

```python
column_subset = movies_df.loc[:,"genre":"director"]
```

```python
column_subset
```

Output:
```
genre  \
Title                                               
Guardians of the Galaxy   Action,Adventure,Sci-Fi   
Prometheus               Adventure,Mystery,Sci-Fi   
Split                             Horror,Thriller   
Sing                      Animation,Comedy,Family   
Suicide Squad            Action,Adventure,Fantasy   
...                                           ...   
Secret in Their Eyes          Crime,Drama,Mystery   
Hostel: Part II                            Horror   
Step Up 2: The Streets        Drama,Music,Romance   
Search Party                     Adventure,Comedy   
Nine Lives                  Comedy,Family,Fantasy   

                                                               description  \
Title                                                                        
Guardians of the Galaxy  A group of intergalactic criminals are forced ...   
Prometheus               Following clues to the origin of mankind, a te...   
Split                    Three girls are k
```

Notice how we put `:` in the first slot (i.e. before the comma)? This says give us all the rows. Similarly with `.iloc`.

```python
column_subset = movies_df.iloc[:,1]
```

```python
column_subset
```

Output:
```
Title
Guardians of the Galaxy     Action,Adventure,Sci-Fi
Prometheus                 Adventure,Mystery,Sci-Fi
Split                               Horror,Thriller
Sing                        Animation,Comedy,Family
Suicide Squad              Action,Adventure,Fantasy
                                     ...           
Secret in Their Eyes            Crime,Drama,Mystery
Hostel: Part II                              Horror
Step Up 2: The Streets          Drama,Music,Romance
Search Party                       Adventure,Comedy
Nine Lives                    Comedy,Family,Fantasy
Name: genre, Length: 1000, dtype: object
```

```python
column_subset = movies_df.iloc[:,1:4]
```

```python
column_subset
```

Output:
```
genre  \
Title                                               
Guardians of the Galaxy   Action,Adventure,Sci-Fi   
Prometheus               Adventure,Mystery,Sci-Fi   
Split                             Horror,Thriller   
Sing                      Animation,Comedy,Family   
Suicide Squad            Action,Adventure,Fantasy   
...                                           ...   
Secret in Their Eyes          Crime,Drama,Mystery   
Hostel: Part II                            Horror   
Step Up 2: The Streets        Drama,Music,Romance   
Search Party                     Adventure,Comedy   
Nine Lives                  Comedy,Family,Fantasy   

                                                               description  \
Title                                                                        
Guardians of the Galaxy  A group of intergalactic criminals are forced ...   
Prometheus               Following clues to the origin of mankind, a te...   
Split                    Three girls are k
```

We can also access a subset of rows and columns _at the same time_.

```python
data_subset = movies_df.loc["Prometheus":"Sing", "genre":"director"]
```

```python
data_subset
```

Output:
```
genre  \
Title                                  
Prometheus  Adventure,Mystery,Sci-Fi   
Split                Horror,Thriller   
Sing         Animation,Comedy,Family   

                                                  description  \
Title                                                           
Prometheus  Following clues to the origin of mankind, a te...   
Split       Three girls are kidnapped by a man with a diag...   
Sing        In a city of humanoid animals, a hustling thea...   

                        director  
Title                             
Prometheus          Ridley Scott  
Split         M. Night Shyamalan  
Sing        Christophe Lourdelet
```

```python
data_subset = movies_df.iloc[1:4, 1:4]
```

```python
data_subset
```

Output:
```
genre  \
Title                                  
Prometheus  Adventure,Mystery,Sci-Fi   
Split                Horror,Thriller   
Sing         Animation,Comedy,Family   

                                                  description  \
Title                                                           
Prometheus  Following clues to the origin of mankind, a te...   
Split       Three girls are kidnapped by a man with a diag...   
Sing        In a city of humanoid animals, a hustling thea...   

                        director  
Title                             
Prometheus          Ridley Scott  
Split         M. Night Shyamalan  
Sing        Christophe Lourdelet
```

## Filtering Data

Often times, we really want to select data based on some condition, or really, to filter it. We can do that with DataFrames quite easily as well. Suppose we only want movies that were directed by Ridley Scott.

```python
movies_df[movies_df['director'] == "Ridley Scott"]
```

Output:
```
rank                     genre  \
Title                                                    
Prometheus                 2  Adventure,Mystery,Sci-Fi   
The Martian              103    Adventure,Drama,Sci-Fi   
Robin Hood               388    Action,Adventure,Drama   
American Gangster        471     Biography,Crime,Drama   
Exodus: Gods and Kings   517    Action,Adventure,Drama   
The Counselor            522      Crime,Drama,Thriller   
A Good Year              531      Comedy,Drama,Romance   
Body of Lies             738      Action,Drama,Romance   

                                                              description  \
Title                                                                       
Prometheus              Following clues to the origin of mankind, a te...   
The Martian             An astronaut becomes stranded on Mars after hi...   
Robin Hood              In 12th century England, Robin and his band of...   
American Gangster       In 1970s America, a detective work
```

What if we want only movies with ratings greater than or equal to 8.6?

```python
movies_df[movies_df['rating'] >= 8.6]
```

Output:
```
rank                    genre  \
Title                                             
Interstellar        37   Adventure,Drama,Sci-Fi   
The Dark Knight     55       Action,Crime,Drama   
Inception           81  Action,Adventure,Sci-Fi   
Kimi no na wa       97  Animation,Drama,Fantasy   
Dangal             118   Action,Biography,Drama   
The Intouchables   250   Biography,Comedy,Drama   

                                                        description  \
Title                                                                 
Interstellar      A team of explorers travel through a wormhole ...   
The Dark Knight   When the menace known as the Joker wreaks havo...   
Inception         A thief, who steals corporate secrets through ...   
Kimi no na wa     Two strangers find themselves linked in a biza...   
Dangal            Former wrestler Mahavir Singh Phogat and his t...   
The Intouchables  After he becomes a quadriplegic from a paragli...   

                           director  \
T
```

Often times we really want to filter by more than one criterion. For example, we might want to get movies that are directed either by Christopher Nolan **or** Ridley Scott. To do an **or** we just connect the conditions by `|`, like below.

```python
movies_df[(movies_df['director'] == 'Christopher Nolan') | (movies_df['director'] == 'Ridley Scott')]
```

Output:
```
rank                     genre  \
Title                                                    
Prometheus                 2  Adventure,Mystery,Sci-Fi   
Interstellar              37    Adventure,Drama,Sci-Fi   
The Dark Knight           55        Action,Crime,Drama   
The Prestige              65      Drama,Mystery,Sci-Fi   
Inception                 81   Action,Adventure,Sci-Fi   
The Martian              103    Adventure,Drama,Sci-Fi   
The Dark Knight Rises    125           Action,Thriller   
Robin Hood               388    Action,Adventure,Drama   
American Gangster        471     Biography,Crime,Drama   
Exodus: Gods and Kings   517    Action,Adventure,Drama   
The Counselor            522      Crime,Drama,Thriller   
A Good Year              531      Comedy,Drama,Romance   
Body of Lies             738      Action,Drama,Romance   

                                                              description  \
Title                                                                       
```

Given that we wanted to check whether or not a value was in a set of things (in this case whether or not "director" is in `["Christopher Nolan", "Ridley Scott"]`), we could have done something simpler and used the `.isin()` method.

```python
movies_df[movies_df['director'].isin(['Christopher Nolan', 'Ridley Scott'])]
```

Output:
```
rank                     genre  \
Title                                                    
Prometheus                 2  Adventure,Mystery,Sci-Fi   
Interstellar              37    Adventure,Drama,Sci-Fi   
The Dark Knight           55        Action,Crime,Drama   
The Prestige              65      Drama,Mystery,Sci-Fi   
Inception                 81   Action,Adventure,Sci-Fi   
The Martian              103    Adventure,Drama,Sci-Fi   
The Dark Knight Rises    125           Action,Thriller   
Robin Hood               388    Action,Adventure,Drama   
American Gangster        471     Biography,Crime,Drama   
Exodus: Gods and Kings   517    Action,Adventure,Drama   
The Counselor            522      Crime,Drama,Thriller   
A Good Year              531      Comedy,Drama,Romance   
Body of Lies             738      Action,Drama,Romance   

                                                              description  \
Title                                                                       
```

However, often times, we want to combine not just **or** statements or checking whether values are **in** something, but we want to chain together a bunch of filters and check if they all hold. That is an **and** statement and we use `&` to get it. Suppose that we wanted movies released between 2005 and 2010 with a rating above 8.0 but made below the 25th percentile in revenue. We can do this with the below.

```python
movies_df[((movies_df['year'] >= 2005) & (movies_df['year'] <= 2010))
          & (movies_df['rating'] > 8.0)
          & (movies_df['revenue_millions'] < movies_df['revenue_millions'].quantile(0.25))]
```

Output:
```
rank               genre  \
Title                                           
3 Idiots              431        Comedy,Drama   
The Lives of Others   477      Drama,Thriller   
Incendies             714   Drama,Mystery,War   
Taare Zameen Par      992  Drama,Family,Music   

                                                           description  \
Title                                                                    
3 Idiots             Two friends are searching for their long lost ...   
The Lives of Others  In 1984 East Berlin, an agent of the secret po...   
Incendies            Twins journey to the Middle East to discover t...   
Taare Zameen Par     An eight-year-old boy is thought to be a lazy ...   

                                             director  \
Title                                                   
3 Idiots                              Rajkumar Hirani   
The Lives of Others  Florian Henckel von Donnersmarck   
Incendies                            Denis Villeneuve
```

### Problem #3 - 1 point

Create a new dataframe called `chipo_filtered` that consists of all rows where more than one `Chicken Bowl` with the same `choice_description` was ordered. Note that each row corresponds to a particular item with a particular choice description. The `quantity` column denotes how many of that item with that choice description was in the order.

The grade cell will check your answer. If you do not get any errors, you very likely have gotten this question correct. However, there are held back tests as well.

<!-- grade_id: problem_3, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
file_path = 'chipotle.tsv'
chipo_df = pd.read_csv(file_path, delimiter='\t')
chicken_bowl_orders = chipo_df[chipo_df['item_name'] == 'Chicken Bowl']
grouped = chicken_bowl_orders.groupby('choice_description')['quantity'].sum().reset_index()
filtered_groups = grouped[grouped['quantity'] > 1]
chipo_filtered = pd.merge(chipo_df, filtered_groups, on='choice_description')
print(chipo_filtered.head())
```

Output:
```
order_id  quantity_x      item_name  \
0         2           2   Chicken Bowl   
1       484           1     Steak Bowl   
2       822           1     Steak Bowl   
3         3           1   Chicken Bowl   
4       519           1  Steak Burrito   

                                  choice_description item_price  quantity_y  
0  [Tomatillo-Red Chili Salsa (Hot), [Black Beans...    $16.98            2  
1  [Tomatillo-Red Chili Salsa (Hot), [Black Beans...     $8.99            2  
2  [Tomatillo-Red Chili Salsa (Hot), [Black Beans...     $8.99            2  
3  [Fresh Tomato Salsa (Mild), [Rice, Cheese, Sou...    $10.98            2  
4  [Fresh Tomato Salsa (Mild), [Rice, Cheese, Sou...    $11.48            2
```

<!-- grade_id: problem_3_test, points: 1, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT AND DO NOT COPY.
# This will give an error if your answer is wrong.
from nose.tools import assert_equal, assert_in
assert_equal(len(chipo_filtered), 33)
assert_equal(chipo_filtered.iloc[3]['quantity'], 3)
assert_equal(chipo_filtered.iloc[24]['order_id'], 1374)
```

Output:
```
File /opt/conda/lib/python3.11/unittest/case.py:873, in TestCase.assertEqual(self, first, second, msg)
    869 """Fail if the two objects are unequal as determined by the '=='
    870    operator.
    871 """
    872 assertion_func = self._getAssertEqualityFunc(first, second)
--> 873 assertion_func(first, second, msg=msg)

File /opt/conda/lib/python3.11/unittest/case.py:866, in TestCase._baseAssertEqual(self, first, second, msg)
    864 standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    865 msg = self._formatMessage(msg, standardMsg)
--> 866 raise self.failureException(msg)

AssertionError: 1085 != 33
```

## Using `.groupby` to Answer Questions About Groups

A very useful method on dataframes is the `.groupby()` method. `.groupby()` allows you to **group** the data **by** a certain column. This is very similar to a pivot table in excel. Once you have grouped the data, you can do useful operations like find the minimum rating over all movies for a particular directory. First, you need to group the data.

```python
movies_grouped = movies_df.groupby('director')
```

We can now easily answer questions like, how many directors are in the data set. That is just the length of the `movies_grouped` table, which we can get by this.

```python
len(movies_grouped)
```

Output:
```
644
```

Once we have the movies grouped together by `director`, we can pull out individual directors and look at the movies they have worked on.

```python
movies_grouped.get_group('Zack Snyder')
```

Output:
```
rank                     genre  \
Title                                                                
Batman v Superman: Dawn of Justice    61   Action,Adventure,Sci-Fi   
300                                  114        Action,Fantasy,War   
Watchmen                             148      Action,Drama,Mystery   
Sucker Punch                         286            Action,Fantasy   
Man of Steel                         295  Action,Adventure,Fantasy   

                                                                          description  \
Title                                                                                   
Batman v Superman: Dawn of Justice  Fearing that the actions of Superman are left ...   
300                                 King Leonidas of Sparta and a force of 300 men...   
Watchmen                            In 1985 where former superheroes exist, the mu...   
Sucker Punch                        A young girl is institutionalized by her abusi...   
Man of Stee
```

We can also use methods like `.sum()`, `.min()`, `.max()`, `.median`, `.count()`, and others to get a look into the data. Below we sum up all "summable" columns by each director.

```python
movies_grouped.sum()
```

Output:
```
rank                                              genre  \
director                                                                       
Aamir Khan            992                                 Drama,Family,Music   
Abdellatif Kechiche   312                                      Drama,Romance   
Adam Leon             784                                     Comedy,Romance   
Adam McKay           1910  Biography,Comedy,DramaComedyAction,Comedy,Crim...   
Adam Shankman        1460            Comedy,Drama,FamilyComedy,Drama,Musical   
...                   ...                                                ...   
Xavier Dolan         1588                                         DramaDrama   
Yimou Zhang             6                           Action,Adventure,Fantasy   
Yorgos Lanthimos      479                 Comedy,Drama,RomanceDrama,Thriller   
Zack Snyder           904  Action,Adventure,Sci-FiAction,Fantasy,WarActio...   
Zackary Adler         605                                    
```

We can then pull out individual directors and see their particular stats.

```python
movies_grouped.sum().loc['Zack Snyder']
```

Output:
```
rank                                                              904
genre               Action,Adventure,Sci-FiAction,Fantasy,WarActio...
description         Fearing that the actions of Superman are left ...
actors              Ben Affleck, Henry Cavill, Amy Adams, Jesse Ei...
year                                                            10055
runtime                                                           683
rating                                                           35.2
votes                                                         2301544
revenue_millions                                               975.74
metascore                                                       240.0
Name: Zack Snyder, dtype: object
```

`.sum()` is probably not very useful here. However, we may want to know what Zack Snyder's best rated movie was. We can do that with `.max()`.

```python
movies_grouped.max().loc['Zack Snyder']
```

Output:
```
rank                                                              295
genre                                              Action,Fantasy,War
description         King Leonidas of Sparta and a force of 300 men...
actors              Jackie Earle Haley, Patrick Wilson, Carla Gugi...
year                                                             2016
runtime                                                           162
rating                                                            7.7
votes                                                          637104
revenue_millions                                               330.25
metascore                                                        56.0
Name: Zack Snyder, dtype: object
```

You want to be a little careful about interpreting exactly what is going on above. the `.max()` operator is taking the maximum over all columns independently, so Zack Snyder's highest ranked film had a metascore of 56, but the description above is not about this movie. When we looked at all of Zack Snyder's films above, we can see that this metascore was associated with "Watchmen", but the description is of the movie "300".

Often times, we really need to sort the values to answer these questions.

## Sorting Values while Grouping

Dataframes also make sorting easy. Suppose we wanted to quickly answer the question about which of Zack Snyder's movies had the highest metascore, we can first get the group of all of Zack Snyder's movies (like above), and then we can sort by metascore.

```python
movies_grouped.get_group('Zack Snyder').sort_values(by=['metascore'], ascending=False)
```

Output:
```
rank                     genre  \
Title                                                                
Watchmen                             148      Action,Drama,Mystery   
Man of Steel                         295  Action,Adventure,Fantasy   
300                                  114        Action,Fantasy,War   
Batman v Superman: Dawn of Justice    61   Action,Adventure,Sci-Fi   
Sucker Punch                         286            Action,Fantasy   

                                                                          description  \
Title                                                                                   
Watchmen                            In 1985 where former superheroes exist, the mu...   
Man of Steel                        Clark Kent, one of the last of an extinguished...   
300                                 King Leonidas of Sparta and a force of 300 men...   
Batman v Superman: Dawn of Justice  Fearing that the actions of Superman are left ...   
Sucker Punc
```

What we did above is we sorted values using `.sort_values()` first `by=['metascore']` and we told it that we wanted values sorted in descending order by setting `ascending=False` (ascending order is the default).

We can combine multiple of these operations to really start to get a good sense for our data. Suppose we wanted to know which director had the highest average metascore, we could use the `.mean()` method on the groups to get the mean of all of the columns by group, and then we can sort by metascore. Again, note the use of `numeric_only=True`.

```python
movies_grouped.mean(numeric_only=True).sort_values(by=['metascore'], ascending=False)
```

Output:
```
rank    year  runtime  rating     votes  revenue_millions  \
director                                                                       
Barry Jenkins      42.0  2016.0    111.0     7.5  135095.0             27.85   
Kenneth Lonergan   22.0  2016.0    137.0     7.9  134213.0             47.70   
Todd Haynes       502.0  2015.0    118.0     7.2   77995.0              0.25   
Kathryn Bigelow   540.0  2010.0    144.0     7.5  289342.0             55.71   
Michael Goi       634.0  2011.0     85.0     4.9    6683.0               NaN   
...                 ...     ...      ...     ...       ...               ...   
S.S. Rajamouli     27.0  2015.0    159.0     8.3   76193.0              6.50   
Saul Dibb         821.0  2014.0    107.0     6.9   13711.0               NaN   
Shawn Burkett      43.0  2016.0     73.0     2.7     496.0               NaN   
Todor Chapkanov   124.0  2016.0     86.0     7.4   10428.0               NaN   
Warren Beatty     748.0  2016.0    127.0     5.7    3731.0 
```

It turns out Barry Jenkins has the highest average metascore. Let's see what he made.

```python
movies_grouped.get_group('Barry Jenkins')
```

Output:
```
rank  genre                                        description  \
Title                                                                       
Moonlight    42  Drama  A chronicle of the childhood, adolescence and ...   

                director                                             actors  \
Title                                                                         
Moonlight  Barry Jenkins  Mahershala Ali, Shariff Earp, Duan Sanderson, ...   

           year  runtime  rating   votes  revenue_millions  metascore  
Title                                                                  
Moonlight  2016      111     7.5  135095             27.85       99.0
```

Now, let's see which director has made the most movies.

```python
movies_grouped.count().sort_values(by=['year'], ascending=False)
```

Output:
```
rank  genre  description  actors  year  runtime  rating  \
director                                                                      
Ridley Scott           8      8            8       8     8        8       8   
M. Night Shyamalan     6      6            6       6     6        6       6   
David Yates            6      6            6       6     6        6       6   
Michael Bay            6      6            6       6     6        6       6   
Paul W.S. Anderson     6      6            6       6     6        6       6   
...                  ...    ...          ...     ...   ...      ...     ...   
Ilya Naishuller        1      1            1       1     1        1       1   
Ido Fluk               1      1            1       1     1        1       1   
Hugo Gélin             1      1            1       1     1        1       1   
Hope Dickson Leach     1      1            1       1     1        1       1   
Zackary Adler          1      1            1       1     1        1     
```

In this dataset, Ridley Scott has made the most movies. Notice we sorted by year because count just counts the values in each column that aren't missing and we just needed something without missing values.

### Problem #4 - 1 point

What is the most ordered item in the dataset? You likely will want to use both `.groupby()` and `.sum()` here. The grade cell will not tell you whether or not you are right. Your answer should look like
``` python
most_ordered = "Item Name Typed Exactly"
```
With `Item Name Typed Exactly` replaced with the exact item name in quotes (and not indented obviously). I.e., if the most common item was `Chips and Fresh Tomato Salsa`, then your answer would look like `most_ordered = "Chips and Fresh Tomato Salsa"`.

<!-- grade_id: problem_4, points: 0, grade: false, solution: true -->
```python
import pandas as pd
file_path = 'chipotle.tsv'
chipo_df = pd.read_csv(file_path, delimiter='\t')
items_grouped = chipo_df.groupby('item_name')['quantity'].sum()
most_ordered = items_grouped.idxmax()
most_ordered
```

Output:
```
'Chicken Bowl'
```

<!-- grade_id: problem_4_test, points: 1, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
# This cell will not tell you whether or not your answer is right.
# If you have not assigned something to the variable most_ordered or you have not assigned a string,
# you will get an error, but if you just get the question wrong,
# you will not necessarily get an error.
from nose.tools import assert_equal, assert_in
assert_equal(type(most_ordered), str)
print(most_ordered)
```

Output:
```
Chicken Bowl
```

### Problem #5 - 1 point

What is the average revenue amount per order? You should assign your answers to the variable `average_revenue`. You can round to the second decimal place (i.e. `10.11` is okay, as is `10.11432` but not `10.1`). The grade cell will not tell you whether or not you are right. Your answer should look something like the below:
``` python
average_revenue = x
```
with x replaced with your answer as a number (and not indented, obviously).

<!-- grade_id: problem_5, points: 0, grade: false, solution: true -->
```python
import pandas as pd
file_path = 'chipotle.tsv'
chipo_df = pd.read_csv(file_path, delimiter='\t')
chipo_df['item_price'] = chipo_df['item_price'].str.replace('$', '').astype(float)
order_revenue = chipo_df.groupby('order_id')['item_price'].sum()
average_revenue = order_revenue.mean()
average_revenue = round(average_revenue, 2)
print(average_revenue)
```

Output:
```
18.81
```

<!-- grade_id: problem_5_test, points: 1, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
# This cell will not tell you whether or not your answer is right.
# If you have not assigned something to the variable average_revenue or you have not assigned a number,
# you will get an error, but if you just get the question wrong,
# you will not necessarily get an error.
from nose.tools import assert_equal, assert_true
import numbers
assert_true(isinstance(average_revenue, numbers.Number))
print(average_revenue)
```

Output:
```
18.81
```